/*
 * Copyright (c) 2018, the Dart project authors.  Please see the AUTHORS file
 * for details. All rights reserved. Use of this source code is governed by a
 * BSD-style license that can be found in the LICENSE file.
 */
/**
 * @assertion We say that a type T0 is a subtype of a type T1 (written T0 <: T1)
 * when:
 * ype Variable Reflexivity 1: T0 is a type variable X0 or a promoted type
 * variables X0 & S0 and T1 is X0.
 * @description Check that if a type T0 is a promoted type variables X0 & S0
 * and T1 is another subtype X0, then a type T0 is not a subtype of a type T1.
 * @author ngl@unipro.ru
 */

import "../utils/common.dart";
import "../../../Utils/expect.dart";


class X0 {}
class S0 extends X0 {}
class C<T> {}
class S1 extends X0 {}


C<S0> t0Instance = new C<S0>();
C<S1> t1Instance = new C<S1>();





C<S1> returnValueFunc() => forgetType(t0Instance);

class ReturnValueTest {
  static C<S1> staticTestMethod() => forgetType(t0Instance);

  C<S1> testMethod() => forgetType(t0Instance);

  C<S1> get testGetter => forgetType(t0Instance);
}

class ReturnValueGen<X> {
  X testMethod() => forgetType(t0Instance);
  X get testGetter => forgetType(t0Instance);
}

main() {
  C<S1> returnValueLocalFunc() => forgetType(t0Instance);

  Expect.throws(() {returnValueFunc();}, (e) => e is TypeError);
  Expect.throws(() {returnValueLocalFunc();}, (e) => e is TypeError);
  Expect.throws(() {ReturnValueTest.staticTestMethod();}, (e) => e is TypeError);

  Expect.throws(() {new ReturnValueTest().testMethod();}, (e) => e is TypeError);
  Expect.throws(() {new ReturnValueTest().testGetter;}, (e) => e is TypeError);

  Expect.throws(() {new ReturnValueGen<C<S1>>().testMethod();}, (e) => e is TypeError);
  Expect.throws(() {new ReturnValueGen<C<S1>>().testGetter;}, (e) => e is TypeError);
}
